```python
# Initialize Otter
import otter
grader = otter.Notebook("hw10.ipynb")
```

# Homework 10: Pandas

**Recommended Readings**: 

* [Working with Dataframes using pandas](https://learningds.org/ch/06/pandas_intro.html)

**Helpful Resources:**

- [Documentation for the datascience module](http://data8.org/datascience/index.html)
- [Pandas documentation](https://pandas.pydata.org/pandas-docs/stable/dsintro.html)

Please complete this notebook by filling in the cells provided. 

**Note: As usual, this homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged.


---

[Pandas](https://pandas.pydata.org/) is one of the most widely used Python libraries in data science. In this assignment, you will review commonly used data-wrangling operations/tools in `pandas`. We aim to give you familiarity with:

* Creating `DataFrames`,
* Slicing `DataFrames` (i.e., selecting rows and columns)
* Filtering data (using boolean arrays)

In this assignment, you are going to use several `pandas` methods. 

`pandas` is very similar to the `datascience` library that we have been using thus far. The first part of this assignment is a sort of tutorial that will guide you through the similarities and differences.

**Note**: The `pandas` interface is notoriously confusing for beginners, and the documentation is not consistently great. Throughout the semester, you will have to search through `pandas` documentation and experiment, but remember it is part of the learning experience and will help shape you as a data scientist!

**This assignment seems long, but rest assured that a vast majority of it is a tutorial (i.e., we will guide you through many aspects of using `pandas` in the most efficient way possible!).**

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
from datascience import *
%matplotlib inline
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part I: From `datascience` library to Python's `pandas` library

**Tip: None of Part I is graded. It is a tutorial. If you're in a rush, go to Part II first and come back to Part I if you need help.**

Throughout this course, we have been working a lot with the `datascience` library, a library created by faculty at UC Berkeley, specifically for courses like this. While this library is not used outside of this course, all of the ideas and concepts behind the library and the different functions are definitely used when dealing with data science problems in the real world. One of the common libraries used in industry is called `pandas`, and is a way to structure and analyze rectangular/tabular data. Using the `datascience` library in DASC130 is a solid stepping stone to understanding `pandas` better. Throughout this notebook, we will be going over certain concepts that we saw in the `datascience` library and showing the equivalent functions that we will use in `pandas`. The syntax and function names may be different but the underlying concepts are still the same! 

Above, we `import pandas as pd`, which means that any function associated with pandas should be called using `pd.function_name()`. This tells Python that we want to use the specific `function_name` from the Pandas library. In theory, we could do `import pandas as pandas` or any other name, but it is known and commonly used to import it as `pd`. We will see some examples of this later in this notebook.


---

### Tables and DataFrames

In `datascience`, we have something called a Table, which is a way to organize your data in a tabular format, which makes accessing rows and columns of data easier. In `pandas`, this structure is called a DataFrame. A DataFrame is the primary data structure in `pandas`. Similar to a Table, we can access different rows and columns of a DataFrame. Tables are essentially the same as DataFrames: we can do similar actions and functions with both.

In the following lines of code, we create a Table and DataFrame of data by importing and reading an external [csv file](https://en.wikipedia.org/wiki/Comma-separated_values). The Table will be called `cones_table` and the DataFrame will be called `cones_df`.

```python
# datascience
cones_table = Table.read_table('cones.csv') # read the cones csv file into the notebook
```

```python
# pandas
cones_df = pd.read_csv('cones.csv') # read the cones csv file into the notebook
```

```python
cones_table
```

```python
cones_df
```

We see here how to import an external csv file into our current notebook using both libraries. With `pandas` `read_csv()` function, we use `pd.read_csv()` to let Python know we want to use this particular `pandas` function.
Notice how there is an extra column of numbers in the DataFrame that results. These numbers are called the index of a DataFrame, and will be useful later in order to select different rows that we are interested in. 

In `pandas`, we have two types of objects: DataFrames and Series. Series are similar to a DataFrame except there is only one column in a Series. You can think of Series as a single column of a DataFrame. Later in this notebook, we will use Series as well as DataFrames. Series come up when we select one column of a DataFrame, apply different functions to a column, etc.

You can learn more about DataFrames and Series with this [documentation](https://pandas.pydata.org/pandas-docs/stable/dsintro.html) as well.


---

### Make a Table/DataFrame
If we want to create a new Table or DataFrame from scratch, we can call the respective functions and then specify the columns and the data within each column (where the data is in array/list format). Some examples from both `datascience` and `pandas` are shown below.

```python
# datascience
flowers = Table().with_columns(
    'Number of petals', make_array(8, 34, 5),
    'Name', make_array('lotus', 'sunflower', 'rose')
) # create Table with 2 columns and 3 rows
flowers
```

```python
# pandas
flowers1 = pd.DataFrame(data = {'Number of petals': [8, 34, 5], 'Name': ['lotus', 'sunflower', 'rose']})
flowers1 # create DataFrame with 2 columns and 3 rows
```


---

#### Aside: Python Dictionaries

To create `flowers1`, we used a Python data structure called a [dictionary](https://docs.python.org/3/tutorial/datastructures.html#dictionaries). If you're not already familiar with this data structure, a dictionary is a set of key:value pairs which are instantiated with curly brackets `{}`.

```python
# Below is an empty dictionary
empty_dict = {}
empty_dict
```

```python
# Below is a dictionary with three key-value pairs
my_words_and_definitions = {
    'dragon': 'a mythical creature resembling a giant reptile that can fly and breathe fire',
    'unicorn': 'a mythical horse-like animal with a single horn on its forehead',
    'phoenix': 'a mythical bird that is said to burst into flames and be reborn from its ashes'
}
# 'unicorn' is the key, and its definition is the value
# To access the definition of 'unicorn', we use the key 'unicorn'
my_words_and_definitions['unicorn']
```

```python
# As we saw above, lists can also be used as values in a dictionary:
my_favorite_things = {
    'colors': ['blue', 'green', 'purple'],
    'numbers': [7, 42, 3.14],
    'foods': ['pizza', 'sushi', 'chocolate']
}
# To access the list of favorite colors, we use the key 'colors'
my_favorite_things['colors']
```


---

### Select Columns
In both `datascience` and `pandas` we can select a column or multiple columns, based on what information we want from the data. In `datascience` we use the `select()` function, while in `pandas` we use indexing.

```python
# datascience
# select one column
cones_table.select("Color")
```

```python
# datascience
# select multiple columns
cones_table.select("Flavor", "Color")
```

```python
# pandas
# select one column - return Series
cones_df["Color"]
```

```python
# pandas
# select one column - return DataFrame
cones_df[["Color"]]
```

```python
# pandas
# select a specific column
cones_df.Color
```

```python
# pandas
# select multiple columns
cones_df[["Flavor", "Color"]]
```

We can also use `loc` and `iloc` to select specific columns, which we will introduce in the next section.


---

### Select Rows
In `datascience`, we use `take` in order to select certain rows, based on what row numbers we want to select (0 indexed). With selecting multiple rows, we use the concept of list slicing in order to select a sequence of rows, or even select multiple rows (if the numbers are in an array). 

```python
# datascience 
cones_table.take(2) # select the row with an index of 2
```

```python
cones_table.take(np.arange(1, 3)) # select the rows with index from 1 to 3 (not 3)
```

```python
cones_table.take[0, 3, 4] # select the rows with index 0, 3, 4
```

```python
# pandas
cones_df[0:3] # select the rows with index from 0 to 3 (not including 3)
```


---

#### loc and iloc

With `pandas`, we can use `loc` and `iloc` to select certain rows or columns from a DataFrame - all at once - making it a powerful tool. `loc` gets rows or columns with particular labels from the index. `iloc` gets rows or columns at particular positions in the index (so it only takes integers). With both of these, we index using brackets and specify which rows and columns we want to select based on `[row, column]`. Note: if we have just `:` for either row or column, this means that we select all of it (based on if it is in the row or column section of the brackets). An example is shown below.

```python
cones_df.loc[:,'Flavor'] # select all rows but only the column Flavor
```

```python
cones_df.loc[:,['Flavor', 'Color']] # select all rows but only the Flavor and Color columns
```

```python
cones_df.iloc[1:3,0:3] # select the rows with index from 1 to 3 (not including 3) and columns in the positions of 0 to 3 (not including 3) 
```


---

### Rename Columns

In the `datascience` package, we can rename a column or multiple columns. We may want to do this when we update a column to make the title more specific. We also need to make sure that the new column name is not the same as any other existing columns (so we do not have two columns with the same name as this will error). We do this with the `relabeled()` function. 

```python
cones_table.relabeled('Price', 'New Price') # rename the Price columns as New Price
```

```python
cones_table.relabeled(['Color', 'Price'], ['Ice Cream Color', 'New Price']) # rename two columns: Color as Ice Cream Color and Price as New Price
```

In `pandas` we can use a similar function called `rename()`.

```python
cones_df.rename(columns={'Flavor': 'New Flavor'}) # rename the Flavor column as New Flavor
```

```python
cones_df.rename(columns={'Flavor': 'New Flavor', 'Color': 'New Color'}) # rename Flavor as New Flavor and Color as New Color
```


---

### Where/Filtering

In `datascience` there is a function called `where` which creates a copy of a table with only the rows that match some condition. To filter out a DataFrame by its contents in `pandas`, we need to use boolean (`True` or `False`) expressions in order to select the rows we want to keep.

```python
# datascience
cones_table.where("Price", are.above(4)) # select/filter rows so the Price is above 4
```

We want to keep the rows that have a price that is greater than 4. We can first create a boolean array, which will assign a `True` or `False` value to each row based on if it satisfies the condition we give it. We can then index this array into our original array and only the values that are `True` will be returned in a DataFrame.


```python
# pandas
boolean_array = cones_df["Price"] > 4 # create a boolean array saying if the Price is greater than 4
boolean_array
```

```python
cones_df[boolean_array] # apply boolean array to DataFrame and filter the rows
```

We can also do this all in one step as shown below.

```python
cones_df[cones_df["Price"] > 4] # filter the DataFrame based on Price greater than 4
```

Sometimes we want to filter by multiple different conditions. In `pandas` we can do this using parentheses and a symbol indicating if we want both conditions to be satisfied (and) or at least one to be satisfied (or). The format is: 

`df[(condition1) & (condition2)]`

or 

`df[(condition1) | (condition2)]`

Note: `&` is different from `and` 
also `|` is different from `or`.

```python
cones_df[(cones_df["Price"] > 4) & (cones_df["Price"] < 5)] # filter the DataFrame based on Price greater than 4 and less than 5
```


---

### Sort, Group, Pivot
#### Sort
We can sort values similarly in Tables and DataFrames - using a function and what column we want to sort by.

```python
# datascience
cones_table.sort('Price') # sort the table by Price column
```

```python
# pandas
cones_df.sort_values('Price') # sort DataFrame by Price column
```

If we want to specify whether we sort in ascending order or descending order. In `datascience` we can do this with the `descending` parameter, and sending it equal to `True` or `False`. In `pandas` we can use the `ascending` parameter and set this equal to `True` or `False`. By default, the column sorts in ascending order.

#### Group 
In both `datascience` and `pandas` we can use group functions that allow us to group records of our data into buckets. You can think of group by as splitting the dataset data into buckets. Then you can call "aggregate" functions (`mean`, `sum`, `max`, `min`, etc) on these buckets to find these values per bucket (which can lead to interesting analysis)!

Let's say that we want to group by flavor of ice cream and see what the total sum is.

```python
# datascience
cones_table.select(['Flavor', 'Price']).group('Flavor', collect=sum) # select the Flavor and Price columns and then group by Flavor and find the sum price per flavor
```

```python
# pandas
cones_df[['Flavor', 'Price']].groupby('Flavor').sum() # select the Flavor and Price columns and then group by Flavor and find the sum price per flavor
```

#### Pivot
We can create pivot tables in both `datascience` and `pandas` using different functions and specifying columns, index, values, and the collect/aggregate function acting on the values. 

```python
# datascience
cones_table.pivot(columns = 'Flavor', rows = 'Color', values = 'Price', collect=sum) # create a pivot table with Flavors and Color and sum prices for corresponding entries
```

```python
# pandas
cones_df.pivot_table(values='Price', index=['Color'], columns=['Flavor'], aggfunc=sum) # create a pivot table with Flavors and Color and sum prices for corresponding entries
```

Note: We have `NaN` as values in our above table because `pandas` cannot find appropriate values for those specific combinations of rows and columns. If we want to replace these values with 0, we can use `fillna(0)` on the resulting pivot table. 

Don't be alarmed by these `NaN` values! This is something to note about real data, often times the data is not cleaned already and null values are very common. They may even be important in your exploration of the data, the number of null values you have and where they occur could be important!


---

### Visualizing Data
For the purpose of representing and showing how to plot different visualizations, we will now create another Table/DataFrame.

```python
normal_data_table = Table().with_columns(['data1', np.random.normal(loc = 1, scale = 2, size = 100), 'data2', np.random.normal(loc = 4, scale = 3, size = 100)])
normal_data_table
```

```python
normal_data_df = pd.DataFrame(data = {'data1': np.random.normal(loc = 1, scale = 2, size = 100), 'data2': np.random.normal(loc = 4, scale = 3, size = 100)})
normal_data_df.head(10)
```


---

#### Histograms
`datascience` histograms: http://data8.org/datascience/_autosummary/datascience.tables.Table.hist.html

`pandas` histograms: https://pandas.pydata.org/pandas-docs/version/0.23/generated/pandas.DataFrame.hist.html

```python
# datascience
new = normal_data_table.select('data1') # create histogram of data1
new.hist()
plt.show();
```

```python
# pandas
new = normal_data_df[["data1"]] # create histogram of data1
new.hist()
plt.show();
```


#### Scatterplots
`datascience` scatterplots: http://data8.org/datascience/_autosummary/datascience.tables.Table.scatter.html

`pandas` scatterplots: https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.plot.scatter.html

```python
# datascience
normal_data_table.scatter('data1') # create scatterplot of data1 vs data2
plt.show();
```

```python
# pandas
normal_data_df.plot.scatter(x = 'data1', y = 'data2'); # create scatterplot of data1 vs data2
plt.show();
```

#### Barplots
`datascience` barplots: http://data8.org/datascience/_autosummary/datascience.tables.Table.barh.html

`pandas` barplots: https://pandas.pydata.org/pandas-docs/version/0.23/generated/pandas.DataFrame.plot.bar.html

```python
# datascience
t = Table().with_columns('Furniture', make_array('chairs', 'tables', 'desks'), 'Count', make_array(6, 1, 2), 'Price', make_array(10, 20, 30)) # barplot of count and price of different furnitures
t.barh('Furniture')
plt.show();
```

```python
# pandas
df = pd.DataFrame(data = {'Furniture': ['chairs', 'tables', 'desks'], 'Count': [6, 1, 2], 'Price': [10, 20, 30]})
df.plot.barh('Furniture'); # barplot of count and price of different furnitures
plt.show();
```

For all of the plots above, there are many different parameters you can play around with (color, size, orientation, figure size, axis labels, title, etc). Above are the basic implementations of these graphs but we recommend looking through the documentation and seeing how you can change different aspects of it. Try using the DataFrame given and plotting a different graph!


---

### Export to CSV

When we want to convert a Table to a csv file, we first need to convert it to a DataFrame and then to a csv file. In `pandas`, we can get rid of this intermediate step because we already have a DataFrame!

```python
cones_table.to_df().to_csv('cones_datascience.csv', index = False)
```

```python
cones_df.to_csv('cones_dataframe.csv', index = False)
```


---

### Calculate the number of columns
In `datascience`, calculating the number of columns (or rows) requires a simple call using our Table, either `num_columns` or `num_rows`. With `pandas`, there are two ways to get the number of columns and rows. The `len` function can be used on a part of the DataFrame, or the `shape` function can be used. The `shape` function returns both the number of rows and the number of columns so based on what we want, we will have to select it using indexing. 

```python
# datascience
cones_table.num_columns
```

```python
# pandas
len(cones_df.columns) # find the length of the list of column names
# or
cones_df.shape[1] # select the column part of shape
```


---

### Calculate the number of rows

```python
# datascience
cones_table.num_rows
```

```python
# pandas
len(cones_df) # number of rows in the DataFrame
# or
cones_df.shape[0] # select the row part of shape
```


---

### Apply
In both libraries, we have a function called `apply`, which we use to apply a function on a certain column and all its elements. Note: in `pandas`, apply works on a Series (since this is essentially a single column of a DataFrame).

```python
# datascience
cones_table.apply(round, "Price") # apply the round function to the column Price
```

```python
# pandas
cones_df["Price"].apply(round, "Price") # apply the round function to the column Price
```

Above we see that we get the rounded up prices returned to us. If we wanted to change the column in the original Table/DataFrame, we would set the expression above equal to the column/series so that the change is made and saved in the original.

For example:

`
cones_table = cones_table.apply(round, "Price")
`

or

`
cones_df["Price"] = cones_df["Price"].apply(round, "Price")
`


---

### Joining
Joins are useful when we want to combine two or more tables together, so we can do analysis on all of the tables. In `datascience`, we can use `join` and in `pandas` we can use `merge`. With both we need the information of the two tables and how we are joining on them (the appropriate columns). 

```python
# datascience
table = Table().with_columns('first', make_array('i', 'c', 'c', 'a'), 'second', make_array('a', 'b', 'b', 'j'), 'third', make_array('c', 'd', 'e', 'f'))
table2 = Table().with_columns( 'another', make_array('i', 'a', 'a', 'a'), 'fourth', make_array('a', 'b', 'b', 'j'), 'fifth', make_array('c', 'd', 'e', 'f'))
print(table)
print()
print(table2)
table.join('first', table2, 'fourth') # join table and table 2 together so columns first and fourth match values
```

In the above code block, we create two different tables. When we call the `join()` function on these two tables, we take the cross product of the rows of both tables (every combination of rows that could happen between both tables) and then filter this out based on the columns we specify that need to be the same. In this case we have row `first` from the first table and row `fourth` from the second table, therefore we can only keep rows that have a value in `first` from the first table and `fourth` from the second table that are the same. In our example, the similar value is `first` and `fourth` is the value `a` (we see it in the last row of the `first` column and the first row of the `fourth` column). We then look at the values in these two rows and create a table with the values. Similarly, we can replicate this in `pandas`:

```python
df = pd.DataFrame(data = {'first': ['i', 'c', 'c', 'a'], 'second': ['a', 'b', 'b', 'j'], 'third': ['c', 'd', 'e', 'f']})
df2 = pd.DataFrame(data = {'another': ['i', 'a', 'a', 'a'], 'fourth': ['a', 'b', 'b', 'j'], 'fifth': ['c', 'd', 'e', 'f']})
print(df)
print()
print(df2)
df.merge(df2, left_on = 'first', right_on = 'fourth', how = "inner") # merge df and df2 using an inner join using columns first and fourth
```

In the `pandas` code above, we have an extra argument that we need to specify: how. There are many types of joins that we can do with data. In `datascience`, the only option we can do is an inner join, but in `pandas`, we have the option to do inner, left, right, outer joins. For more information about these types of joins, check out this page: http://pandas.pydata.org/pandas-docs/version/0.19.1/generated/pandas.DataFrame.merge.html


---

### Reading `pandas` Documentation
There are many more functions and methods you can call with `pandas` to do more cool things with Series and DataFrames. One way to learn more about this is by looking through the `pandas` documentation. The documentation has all the different functions associated with `pandas`, and descriptions about what they do, how you use them, and some examples. 

Some tips for reading through the documentation:
* For various functions there are LOTS of different parameters that you can call, usually there are only a few that are important (usually related to the data you are working with and specifying how to run the functions). There are some parameters that are optional and you do not have to specify (automatically Pandas will use default settings for these functions). For example, for joins, an inner join is the default setting but if needed, you can say you want to do an outer join etc.
* Another useful parameter is the `inplace` parameter. When set to true, the data is renamed/the function is run in place (a copy of the new data with the function applied is NOT created)
* For quick lookup for a specific function: in Jupyter you can put your cursor on a Pandas function and press Shift + Tab in order to see part of the documentation. Based on how many times you press the Tab, it will show you the header of the function (how you would call it and parameters) and a description of the function overall and a description of the parameters.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part II: Checking your knowledge

In Part II, do *not* use anything from the `datascience` module!

---

### **REVIEW:** Creating `DataFrames` & Basic Manipulations

Recall that a [DataFrame](https://pandas.pydata.org/pandas-docs/stable/user_guide/dsintro.html#dataframe) is a table in which each column has a specific data type; there is an index over the columns (typically string labels) and an index over the rows (typically ordinal numbers).

Usually, you'll create `DataFrames` by using a function like `pd.read_csv`. However, in this section, we'll discuss how to create them from scratch.

The [documentation](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.html) for the `pandas` `DataFrame` class provides several constructors for the `DataFrame` class.

**Syntax 1:** You can create a `DataFrame` by specifying the columns and values using a dictionary, as shown below. 

The keys of the dictionary are the column names, and the values of the dictionary are lists containing the row entries.

```python
fruit_info = pd.DataFrame(
    data = {'fruit': ['apple', 'orange', 'banana', 'raspberry'],
          'color': ['red', 'orange', 'yellow', 'pink'],
          'price': [1.0, 0.75, 0.35, 0.05]
          })
fruit_info
```

**Syntax 2:** You can also define a `DataFrame` by specifying the rows as shown below. 

Each row corresponds to a distinct tuple, and the columns are specified separately.

```python
fruit_info2 = pd.DataFrame(
    [("red", "apple", 1.0), ("orange", "orange", 0.75), ("yellow", "banana", 0.35),
     ("pink", "raspberry", 0.05)], 
    columns = ["color", "fruit", "price"])
fruit_info2
```

You can obtain the dimensions of a `DataFrame` by using the shape attribute `DataFrame.shape`.

```python
fruit_info.shape
```

You can also convert the entire `DataFrame` into a two-dimensional `NumPy` array. Remember that a `NumPy` array can hold homogenous data whereas a `DataFrame` can contain heterogeneous data. 

```python
numbers = pd.DataFrame({"A":[1, 2, 3], "B":[0, 1, 1]})
numpy_numbers = numbers.to_numpy()

print(type(numpy_numbers))
print(numpy_numbers)
```

The `values` attribute returns the content of the `DataFrame` in the form of a list of lists.

```python
fruit_info.values
```

There are other constructors but we do not discuss them here.


---

### **REVIEW:** Selecting Rows and Columns in `pandas`

As you've seen in lecture and discussion, there are two verbose operators in Python for selecting rows: `loc` and `iloc`. Let's review them briefly.

**Approach 1:** `loc`

The first of the two verbose operators is `loc`, which takes two arguments. The first is one or more **row labels**, the second is one or more **column labels** - both of which are displayed in bold to the left of each of the rows and above each of the columns, respectively. These are not the same as positional indices, which are used for indexing Python lists or `NumPy` arrays!

The desired rows or columns can be provided individually, in slice notation, or as a list. Some examples are given below.

Note that **slicing in `loc` is inclusive** on the provided labels.

```python
# Get rows 0 through 2 (inclusive) with labels 'fruit' through 'price' (which would include the color column that is in between both labels)
fruit_info.loc[0:2, 'fruit':'price']
```

```python
# Get rows 0 through 2 (inclusive) and columns 'fruit' and 'price'. 
# Note the difference in notation and result from the previous example.
fruit_info.loc[0:2, ['fruit', 'price']]
```

```python
# Get rows 0 and 2 and columns fruit and price. 
fruit_info.loc[[0, 2], ['fruit', 'price']]
```

```python
# Get rows 0 and 2 and column fruit
fruit_info.loc[[0, 2], ['fruit']]
```

Note that if we request a single column but don't enclose it in a list, the return type of the `loc` operator is a `Series` rather than a `DataFrame`. 

```python
# Get rows 0 and 2 and column fruit, returning the result as a Series
fruit_info.loc[[0, 2], 'fruit']
```

If we provide only one argument to `loc`, it uses the provided argument to select rows, and returns all columns.

```python
fruit_info.loc[0:1]
```

Note that if you try to access columns without providing rows, `loc` will crash. 

```python
# If you uncomment, this code will crash
# fruit_info.loc[["fruit", "price"]]

# Uncomment, this code works fine: 
# fruit_info.loc[:, ["fruit", "price"]]
```

<br/>

**Approach 2:** `iloc`

`iloc` is very similar to `loc` except that its arguments are **row numbers** and **column numbers**, rather than row and column labels. A useful mnemonic is that the `i` stands for "integer". This is quite similar to indexing into a Python `list` or `NumPy` array.

In addition, **slicing for `iloc` is exclusive** on the provided integer indices. Some examples are given below:

```python
# Get rows 0 through 3 (exclusive) and columns 0 through 3 (exclusive)
fruit_info.iloc[0:3, 0:3]
```

```python
# Get rows 0 through 3 (exclusive) and columns 0 and 2.
fruit_info.iloc[0:3, [0, 2]]
```

```python
# Get rows 0 and 2 and columns 0 and 2.
fruit_info.iloc[[0, 2], [0, 2]]
```

```python
# Get rows 0 and 2 and column fruit
fruit_info.iloc[[0, 2], [0]]
```

```python
# Get rows 0 and 2 and column fruit
fruit_info.iloc[[0, 2], 0]
```

Note that in these `loc` and `iloc` examples above, the row **label** and row **number** were always the same.

Let's see an example where they are different. If we sort our fruits by price, we get:

```python
fruit_info_sorted = fruit_info.sort_values("price")
fruit_info_sorted
```

After sorting, note how row number 0 now has index label 3, row number 1 now has index label 2, etc. These indices are the arbitrary numerical indices generated when we created the `DataFrame`. For example, `banana` was originally in row 2, and so it has row label 2. Note the distinction between the index _label_, and the actual index _position_.

If we request the rows in positions 0 and 2 using `iloc`, we're indexing using the row NUMBERS, not labels. 

```python
fruit_info_sorted.iloc[[0, 2], 0]
```

Lastly, similar to `loc`, the second argument to `iloc` is optional. That is, if you provide only one argument to `iloc`, it treats the argument you provide as a set of desired row numbers, not column numbers.

```python
fruit_info_sorted.iloc[[0, 2]]
```

**Approach 3:** `[]` Notation for Accessing Rows and Columns

`pandas` also supports the `[]` operator. It's similar to `loc` in that it lets you access rows and columns by their name.

However, unlike `loc`, which takes row names and also optionally column names, `[]` is more flexible. If you provide it only row names, it'll give you rows (same behavior as `loc`), and if you provide it with only column names, it'll give you columns (whereas `loc` will crash).

Some examples:

```python
fruit_info[0:2]
```

```python
# Here we're providing a list of fruits as single argument to []
fruit_info[["fruit", "color", "price"]]
```

Note that slicing notation is not supported for columns if you use `[]` notation. Use `loc` instead.

```python
# Uncomment and this code crashes
#fruit_info["fruit":"price"]

# Uncomment and this works fine
#fruit_info.loc[:, "fruit":"price"]
```

`[]` and `loc` are quite similar. For example, the following two pieces of code are functionally equivalent for selecting the fruit and price columns.

1. `fruit_info[["fruit", "price"]]` 
2. `fruit_info.loc[:, ["fruit", "price"]]`.

Because it yields more concise code, you'll find that our code tends to feature `[]`. However, there are some subtle pitfalls of using `[]`. If you're ever having performance issues, weird behavior, or you see a `SettingWithCopyWarning` in `pandas`, switch from `[]` to `loc`, and this may help.

To avoid getting too bogged down in indexing syntax, we'll avoid a more thorough discussion of `[]` and `loc`. We may return to this at a later point in the course.

For more on `[]` vs. `loc`, you may optionally try reading:
1. https://stackoverflow.com/questions/48409128/what-is-the-difference-between-using-loc-and-using-just-square-brackets-to-filte
2. https://stackoverflow.com/questions/38886080/python-pandas-series-why-use-loc/65875826#65875826
3. https://stackoverflow.com/questions/20625582/how-to-deal-with-settingwithcopywarning-in-pandas/53954986#53954986

Now that we've reviewed basic indexing, let's discuss how we can modify `DataFrames`. We'll do this via a series of exercises. 

<br/><br/>

---

### Question 1a

For a `DataFrame` `d`, you can add a column with `d['new column name'] = ...` and assign a `list` or `array` of values to the column. Add a column of integers containing 1, 2, 3, and 4 called `rank1` to the `fruit_info` table, which expresses **your personal preference** about the taste ordering for each fruit (1 is tastiest; 4 is least tasty). There is no right order, it is completely your choice of rankings. **(4 points)**

```python
fruit_info
```

```python
grader.check("q1a")
```

<br/><br/>

---

### Question 1b

You can also add a column to `d` with `d.loc[:, 'new column name'] = ...`. As above, the first parameter is for the rows, and the second is for columns. The `:` means changing all rows, and the `'new column name'` indicates the name of the column you are modifying (or, in this case, adding). 

Add a column called `rank2` to the `fruit_info` table, which contains the same values in the same order as the `rank1` column. **(6 points)**


```python
fruit_info
```

```python
grader.check("q1b")
```

<br/><br/>

---

### Question 2

Use the `.drop()` method to drop both the `rank1` and `rank2` columns you created. Make sure to use the `axis` parameter correctly. Note that `drop` does not change a table but instead returns a new table with fewer columns or rows unless you set the optional `inplace` argument. **(12 points)**

**Hint:** Look through the [documentation](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.drop.html) to see how you can drop multiple columns of a `DataFrame` at once using a list of column names.


```python
fruit_info
```

```python
grader.check("q2")
```

<br/><br/>

---

### Question 3

Use the `.rename()` method to rename the columns of `fruit_info` so they begin with capital letters. Set this new `DataFrame` to `fruit_info_caps`. For an example of how to use rename, see this linked [documentation](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.rename.html). **(12 points)**

```python
fruit_info_caps = ...
fruit_info_caps
```

```python
grader.check("q3")
```

<br/><br/>

---

### Babynames Dataset
For the next few questions of this lab, let's move on to a real-world dataset. We'll be using the babynames dataset from Lecture 3. The babynames dataset contains a record of the given names of babies born in the United States each year.

Let's run the following cells to build the `DataFrame` `babynames`. Note that we only include data from Minnesota due to memory constraints (the full dataset has over 6 million rows!). There should be a total of 125120 records.


```python
file_path = 'MN.TXT'
column_labels = ['State', 'Sex', 'Year', 'Name', 'Count']

babynames = pd.read_csv(file_path, names=column_labels)
```

```python
len(babynames)
```

```python
babynames.head()
```

#### Selection Examples on Baby Names

As with our synthetic fruit dataset, we can use `loc` and `iloc` to select rows and columns of interest from our dataset.

```python
babynames.loc[2:5, 'Name']
```

Notice the difference between the following cell and the previous one; just passing in `'Name'` returns a `Series` while `['Name']` returns a `DataFrame`.

```python
babynames.loc[2:5, ['Name']]
```

The code below collects the rows in positions 1 through 3, and the column in position 3 ("Name").

```python
babynames.iloc[1:4, [3]]
```

<br/><br/>

---

### Question 4a

Use `.loc` to select `Name` and `Year` **in that order** from the `babynames` table. **(8 points)**

```python
name_and_year = ...
name_and_year[:5]
```

```python
grader.check("q4a")
```

<br/><br/>

---

### Question 4b
Now repeat the same selection using the plain `[]` notation. **(8 points)**

```python
name_and_year = ...
name_and_year[:5]
```

```python
grader.check("q4b")
```


---

### **REVIEW**: Filtering with boolean arrays

Filtering is the process of removing unwanted entries. In your quest for cleaner data, you will undoubtedly filter your data at some point: whether it be for clearing up cases with missing values, for culling out fishy outliers, or for analyzing subgroups of your dataset. Example usage looks like `df[df['column name'] < 5]`.

For your reference, some commonly used comparison operators are given below.

Symbol | Usage      | Meaning 
------ | ---------- | -------------------------------------
==   | a == b   | Does a equal b?
<=   | a <= b   | Is a less than or equal to b?
&gt;=   | a >= b   | Is a greater than or equal to b?
<    | a < b    | Is a less than b?
&#62;    | a &#62; b    | Is a greater than b?
~    | ~p       | Returns negation of p
&#124; | p &#124; q | either p OR q
&    | p & q    | both p AND q
^  | p ^ q | p XOR q (exclusive or; only one of p or q)

As an example, in the following, we construct a `DataFrame` containing only rows where the name is Omar.

```python
omar_babynames = babynames[babynames['Name'] == 'Omar']
omar_babynames.head(5)
```

<br/><br/>

---
### Question 5
Using a boolean array, select the names in Year 2000 (from `babynames`) that have larger than 300 counts. Keep all columns from the original `babynames` `DataFrame`.

_Note_: Note that compound expressions have to be grouped with parentheses. That is, any time you use `p & q` to filter the `DataFrame`, make sure to use `df[(df[p]) & (df[q])]` or `df.loc[(df[p]) & (df[q])]`. 

You may use either `[]` or `loc`. Both will achieve the same result. For more on `[]` vs. `loc`, see the stack overflow links from the intro portion of this assignment. **(20 points)**

```python
result = ...
result.head()
```

```python
grader.check("q5")
```


---

### **REVIEW:** `str`

`pandas` provides special purpose functions for working with specific common data types such as strings and dates. For example, the code below provides the length of every baby's name from our `babynames` dataset. 

```python
babynames['Name'].str.len()
```

<br/><br/>

---

### Question 6a

Add a column to `babynames` named `First Letter` that contains the first letter of each baby's name. **(15 points)**

Hint: you can index using `.str` similarly to how you'd normally index Python lists or arrays. Or, you can use `.str.get` [(documentation here)](https://pandas.pydata.org/docs/reference/api/pandas.Series.str.get.html).

```python
babynames['First Letter'] = ...
babynames.head()
```

```python
grader.check("q6a")
```

<br/><br/>

---

### Question 6b

In 2024, how many babies had names that started with the letter "A"? **(15 points)**

```python
babynames_2024 = ...
just_A_names_2024 = ...
number_A_babies = ...
number_A_babies
```

```python
grader.check("q6b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 10!

### Submission Instructions

**Important submission steps:** 
1. **Run** the `grader.check_all` cell below before exporting.
2. Then choose **Save default** from the **File** menu, to save all the outputs.
3. Then **run the final cell**: `grader.export` 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1a (points: None)
```
>>> fruit_info['rank1'].dtype
dtype('int64')
```
```
>>> sorted(fruit_info['rank1'].dropna())
[1, 2, 3, 4]
```


Tests for q1b (points: None)
```
>>> fruit_info['rank2'].dtype
dtype('int64')
```
```
>>> fruit_info['rank2'] == fruit_info['rank1']
0    True
1    True
2    True
3    True
dtype: bool
```


Tests for q2 (points: None)
```
>>> fruit_info_original.shape
(4, 3)
```
```
>>> fruit_info.shape
(4, 5)
```
```
>>> fruit_info_original.columns
Index(['fruit', 'color', 'price'], dtype='object')
```


Tests for q3 (points: None)
```
>>> fruit_info_caps.columns
Index(['Fruit', 'Color', 'Price'], dtype='object')
```
```
>>> fruit_info.columns
Index(['fruit', 'color', 'price', 'rank1', 'rank2'], dtype='object')
```


Tests for q4a (points: None)
```
>>> name_and_year.shape
(407428, 2)
```
```
>>> name_and_year.loc[0, 'Name']
'Mary'
```
```
>>> int(name_and_year.loc[0, 'Year'])
1910
```
```
>>> name_and_year.columns[0]
'Name'
```


Tests for q4b (points: None)
```
>>> name_and_year.shape
(407428, 2)
```
```
>>> name_and_year.loc[0, 'Name']
'Mary'
```
```
>>> int(name_and_year.loc[0, 'Year'])
1910
```
```
>>> name_and_year.columns[0]
'Name'
```


Tests for q5 (points: None)
```
>>> len(result)
10
```
```
>>> int(result['Count'].sum())
35904
```
```
>>> int(result['Count'].iloc[0])
4342
```


Tests for q6a (points: None)
```
>>> 'First Letter' in babynames.columns
True
```
```
>>> babynames['First Letter'].equals(babynames['Name'].str[0])
True
```


Tests for q6b (points: None)
```
>>> years_2022 = babynames_2022['Year'].unique()
>>> assert all(2022 == years_2022)
```
```
>>> assert number_A_names == 55489
```
